Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **34 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Jeong woo Park"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - The Delta Median problem and Randomization

## Question 1 [time estimate: 16 minutes]

### Question 1a [time estimate: 7 minutes]

Write the worst sort function in the world, `worst_sort`. This function takes a list, and then randomly shuffles it until it happens to be in sorted order. Once it is in sorted order then the list is returned.


<!-- grade_id: cell-ead9c74054a1c5da, points: 0, grade: false, solution: true -->
```python
def worst_sort(A):
    
    """
    Sort A in ascending order by randomly shuffling its 
    elements until they are in order.
    
    Parameters
    ----------
    A : list
        A list of numerical values
        
    Returns
    -------
    A : list
        Sorted list
        
    """
    import random
    
    B = A
    while B != sorted(A):
        B = random.shuffle(A)
    
    return A
        
worst_sort([2,1])
```

Output:
```
~\anaconda3\lib\random.py in shuffle(self, x, random)
    304             for i in reversed(range(1, len(x))):
    305                 # pick an element in x[:i+1] with which to exchange x[i]
--> 306                 j = randbelow(i+1)
    307                 x[i], x[j] = x[j], x[i]
    308         else:

~\anaconda3\lib\random.py in _randbelow_with_getrandbits(self, n)
    255         r = getrandbits(k)          # 0 <= r < 2**k
    256         while r >= n:
--> 257             r = getrandbits(k)
    258         return r
    259 

KeyboardInterrupt:
```

<!-- grade_id: cell-ff4db3a3e8b04515, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 1b [time estimate: 3 minutes]
What is the best case complexity of this algorithm?

<!-- grade_id: cell-afac80171b3b6232, points: 0, grade: true, solution: true -->
The algorithm above checks if the input list A is sorted or not. If not sorted, the algorithm gladly takes the list and randomly shuffles it until it becomes a sorted list. So, the best case complexity that we can think of is that the input list A is alredy sorted and thus there is no need to go through the while loop. So, it will be bigTheta(1) time complexity.


### Question 1c [time estimate: 3 minutes]

What is the average case complexity?


<!-- grade_id: cell-9d52d34daa1e3478, points: 0, grade: true, solution: true -->
The average case complexity in this case is a little different. We need to consider the probability of the all possible outcomes. If there are N elements in the list, the list can have N! different outcomes after the shuffle. And thus, after the shuffle the probability of having a perfectly sorted list will be 1/(N!). This case, according to the reading, we need N! complexity to gain a wanted outcome. 

### Question 1d [time estimate: 3 minutes]

For what size lists is this a feasible method?

<!-- grade_id: cell-ab40f08768579798, points: 0, grade: true, solution: true -->
When I assume python can run only a 500,000 unit of time (or else it automatically shuts down) for a jupyter notebook cell it can not have more then 9 elements in the list(9! = 362,880) without considering any other steps that algorithm takes(only considering the permutation calculation).

### [Optional] Question 1e [time estimate: 5 minutes]

Can you think of an even worse sorting method? In such a case, what would its complexity be? How big a list could you sort?

<!-- grade_id: cell-ac05806942341926, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2 [time estimate: 18 minutes]

Approximate median finder. Given a list and $\tilde{\delta}$ (a number between 0 and 0.5), the approximate median finder function returns a number that is guaranteed to lie between the $(50-\tilde{\delta}/2)$% and $(50+\tilde{\delta}/2)$% percentiles, where $\tilde{\delta}=100\cdot \delta$. Implement such a function by randomly selecting an element from the list and testing whether or not it lies within the bounds. If it doesn’t then retry with a new random element, but only a limited number of retries are allowed (you decide the maximum number of retries.)

### Question 2a [time estimate: 5 minutes]
Complete the following function


<!-- grade_id: cell-2ab65512a6148d3c, points: 0, grade: false, solution: true -->
```python
def check_approx_med(A, num, delta):
    """
    Given a list and a number in a list, determine whether it is 
    within (50-delta/2)% and(50+delta/2)% percentiles.
    
    X% percentile of a list is defined to be the smallest number in 
    the list that is as large as at least X% numbers in the list. 
    
    Parameters
    ----------
    A : list
        A list of numerical values
    num 
        The number of interest
    delta : float
        The error, lies between 0 and 0.5

    Returns
    -------
    isApproxMed : bool
        A boolean value indicating whether num is whithin the bound
    
    """
    import random
    import numpy as np
    
    lst = np.array(A)
    
    for i in range(len(A)): #number of retries = length of the given list
        
        #to do with percentile
        #if (num>np.percentile(lst, 0.5-delta/2) and num<np.percentile(lst, 0.5+delta/2)): 
        #but the assertion checking below wants me to calculate the range of numbers from max and min instead of finding the percentile
        mx = max(lst)
        mn = min(lst)
        range_low = (mx-mn)*(0.5-delta/2)
        range_high = (mx-mn)*(0.5+delta/2)
        
        if num>=range_low and num<=range_high:
            return True
        else:
            return False
        
```

<!-- grade_id: cell-df56d7433abb6517, points: 1, grade: true, solution: false -->
```python
assert(check_approx_med([0, 1], 0.4, .25) == True)
assert(check_approx_med([0, 1], 0, .25) == False)
assert(check_approx_med([0], 0, .25) == True)
```

### Question 2b [time estimate: 5 minutes]
Complete the following function that makes use of `check_approx_med` above.


<!-- grade_id: cell-e19685b545147dde, points: 0, grade: false, solution: true -->
```python
def approx_med_finder(A, delta):
    """
    Given a list, find a number in the list that is between 
    50-delta/2% and 50+delta/2% percentiles
    
    Paramters
    ---------
    A : list
        A list of numerical values
    Delta : int
        The error, lies between 0 and 0.5
        
    Returns
    -------
    num
        The approximated median if it is found within 100//delta trials
        (this is one possible maximum number of retries. While you are 
        encouraged to play around with another limits, the submitted work 
        must use this number)
    None
        Finding failed, if nothing is found within 100//delta trials
   
   
    Note
    ----
    The 100//delta is chosen here as it is the expected number of trials for 
    the first successful finding to occur (using geometric distribution)
    
    """
    import random
    i = 1
    while i < 100//delta:
        pick = A[random.randint(0,len(A)-1)]
        if check_approx_med(A, pick, delta) == True:
            return pick
            break
        i += 1
    return None
        
    
import time
approx_med_finder([0,1], 0.5)
time.perf_counter()
```

Output:
```
3453.2670141
```

### Question 2c [time estimate: 3 minutes]

What is the probability of failure in each random trial? What is the probability of failure after all the allowed random trials? Does this scale with $\delta$ or $N$ (the number of elements in a list)?


<!-- grade_id: cell-efe7ab76f6a2546a, points: 0, grade: true, solution: true -->
The smaller the delta becomes, the range of real number of percentile gets smaller and smaller. According to statistics and probability, the smaller the area, the smaller the probbility it becomes for us to have that randum number to be inside the range. Because, the proability of getting the right number that fits the range is delta/1 which is the same thing as a delta.
If the number of elements rise, it does not really matter with statistics that probabability does not change! However, the bigger N, there will be more numbers that will fit in the range too (and also the numbers that doesn't fit inside the range rises too, so the ratio of the numbers that fit inside the range does not change according to size of N)

### Question 2d [time estimate: 5 minutes]
Analyze the expected runtime of `approx_med_finder`. Note that because the function uses `check_approx_med`, you will most likely need to analyze the runtime of that function, too.

<!-- grade_id: cell-86b199798477fc2a, points: 0, grade: true, solution: true -->
Having to find a number that fits the range is delta. This means that it needs to run 1/delta times at least to get one single number that fits the range. So, according to this idea, the total expected runtime of approx_med_finder is (runtime of approx_med_finder)\*(1/delta).